# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #5<span style="color:rgba(0, 76, 153, 1);"><span style="color:red"> $\to$ ($8$ possible marks)</span></span></div>

### Due Thursday, Nov. 2 at 17:00 $-$ upload your .zip submission to the PHYS 231 Canvas Gradebook.
    
<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Operational Ampliers: The Inverting Amplifier

***
## Lab Learning Outcomes:
By the end of the PHYS 231 Lab #5, students will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test a practical op-amp based inverting amplifier.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Experimentally explore the limitations of op-amp circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Use a log-log plot to examine the important characteristics of the inverting amplifier's frequency response.</span></b>

***
## Autograding:
The PHYS 231 Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS231.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS231.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS231.Installer()' command requires the file 'PHYS231.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS231
PHYS231.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 5.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
import uncertainties
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
grader.check("name_and_student_number")
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Oscilloscope Measurement of the Gain</span></div>

## NumPy Arrays

Data that you collect during the lab can be entered directly into a list using the following notation:

```python
x = [x0, x1, x2, x3, ..., xN]
y = [y0, y1, y2, y3, ..., yN]
```
These lists will contain $N+1$ elements. It's usually more convenient to enter data into a NumPy array rather than a list because you can do mathematical operations on arrays that you can't do on lists.  For example:

- squaring an array creates an new array with each of the elements from the original array squared
- multiplying two arrays of the same length creates a new array in which the each elements is the product of the corresponding elements from the original arrays
- ...

NumPy arrays are entered using the following syntax:

```python
x = np.array([x0, x1, x2, x3, ..., xN])
y = np.array([y0, y1, y2, y3, ..., yN])
```             

In the cell below, replace the "..." in NumPy array statements with your measurements of:

- frequency in kHz
- the amplitude of the signal from the function generator Vin in volts
- the uncertainty in Vin in volts
- the amplitude of the output voltage Vout in volts
- the uncertainty in Vout in volts

After executing the cell, the lengths of each of the arrays you've entered will be printed.  Make sure that all of the lengths are the same.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with your data.  Collect enough data that you clearly observe distinct behaviours at low and high frequencies.  We suggest that you plot your data as you collect it.  Please enter your frequency data in **units of kHz**.

*** Please do not change anything to the left of the equals sign. ***

```python
frequency = np.array([..])
Vin = np.array([...])
errVin = np.array([...])
Vout = np.array([...])
errVout = np.array([...])


# The statements below will print the lengths of the arrays.
print("Length of frequency array:", len(frequency))
print("Length of Vin array:", len(Vin))
print("Length of errVin array:", len(errVin))
print("Length of Vout array:", len(Vout))
print("Length of errVout array:", len(errVout))
```

```python
grader.check("q1.1")
```

<!-- END QUESTION -->

## Pandas DataFrames

The Pandas package can be used to display your array data in a nicely-formated table.  The Pandas package creates "DataFrames" which are 2 dimensional data structures, like a 2 dimensional array, or a table with rows and columns.  The cell below creates and then displays a DataFrame using the arrays that you created in the cell above.

```python
# Create the data structure.
data = {
    "frequnecy (kHz)": frequency,
    "Vin (V)": Vin,
    "errVin (V)": errVin,
    "Vout (V)": Vout,
    "errVout (V)": errVout
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(data)
```

## $V_\mathrm{out}/V_\mathrm{in}$ versus $f$

In the cells below, calculate:
  
- $G = V_\mathrm{out}/V_\mathrm{in}$
- the uncertainty in $V_\mathrm{out}/V_\mathrm{in}$

Here, $G$ represents the absolute value of the gain of the amplifier.  We'll assume the relative uncertainty in the frequency is negligible compared the relative uncertainties of the other relevant parameters.

Recall propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(2 marks)</span>**  

In the cell below, replace the ... with expressions that complete the required calculations.

*** Please do not change anything to the left of the equals sign. ***

```python
# Calculate the gain G = Vout/Vin
G = ...

# Calculate the uncertainty in the gain errG.
# Note that, in Python, NumPy can be used to access the square root function.  The syntax is np.sqrt(...).
errG = ...
```

```python
grader.check("q1.2")
```

### Plot the data

The $V_\mathrm{out}/V_\mathrm{in}$ data can now be plotted as a function of frequency.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.3</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate data to generate a plot of $G$ versus frequency.  Be sure to include errror bars.

For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook. 

*** Please do not change anything to the left of the equals sign. ***

```python
PHYS231.Scatter(xData = ..., 
                yData = ..., 
                yErrors = ...,
                xlabel = 'frequency', 
                ylabel = 'Gain', 
                xUnits = 'kHz', 
                yUnits = '');
```

<!-- END QUESTION -->

## log-log Plots

When measuring the response of a system (the gain of the amplifier, in this case) over a wide range of stimulus values (frequency, in this case) it is sometimes convenient to look at the data using *log-log plots*.  In this lab, such a plot will allow us to look for differences in the low-frequency and high-frequency response of the amplifier circuit.

In the cells below, we calculate the log of the the amplifier's gain ($G = V_\mathrm{out}/V_\mathrm{in}$) and the log of frequency.  All you need to do is execute the cells below.

```python
logf = np.log10(frequency)
logG = np.log10(G)
```

To calculate the error in the log of the gain, first note that if:
$$
y= \log x,
$$
Then:
$$
x = 10^y,
$$
such that:
$$
\ln x=\ln\left(10^y\right)=y\ln 10.
$$
Therefore, it must be the case that:
$$
y=\log x=\frac{1}{\ln 10}\ln x,
$$
which allows us to us propagation of errors to deduce that:
$$
\Delta y = \frac{1}{\ln 10}\left(\frac{\Delta x}{x}\right).
$$
In the cell below, enter an expression to calculate the uncertainty in the log of the gain.

***
**<span style="color:blue">Question 1.4</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate expression to calculate the uncertainty in $\log G$. 

*** Please do not change anything to the left of the equals sign. ***

```python
# Note that, in Python, np.log(x) represents ln(x) (natural logarithm, base e) and np.log10(x) represents log(x) (base 10 logarithm).
errlogG = ...
```

```python
grader.check("q1.4")
```

### Plot the $\log G$ vs $log f$

The $\log G$ data can now be plotted as a function of $\log f$.  Execute the cell below to create a scatter plot with error bars.  For additional information about the 'Scatter' function, see **Appendix A** near the end of this notebook. 

```python
PHYS231.Scatter(
    xData = logf,
    yData = logG,
    yErrors = errlogG,
    xlabel = "log f",
    ylabel = "log G",
    xUnits = "",
    yUnits = "",
);
```

## Linear Fit

You should have found that the high-frequency portion of your $\log G$ versus $\log f$ plot is linear with a negative slope.  The final bit of analysis that we'll do is to:

 - select the part of the data that is linear
 - find the slope of the linear section

First, we'll sort the data in the order of increasing values of $\log f$.

```python
inds = logf.argsort()
logf_sorted = logf[inds]
logG_sorted = logG[inds]
errlogG_sorted = errlogG[inds]

# Let's look at the contents of the sorted data in Pandas dataframe.
sortedData = {
    "log f": logf_sorted,
    "log G": logG_sorted,
    "errlog G": errlogG_sorted,
        }

# An option to format the numbers in the DataFrame.
pd.set_option('display.float_format', lambda x: '%0.3g' % x)

#Display the data.
pd.DataFrame(sortedData)
```

Now select the part of the data that will used for the linear fit.  Adjust the start and stop values until the plot shows only the high-frequency linear portion of the data.  You may have to try a few different combinations of start and stop until you obtain the correct portion of linear data.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.5</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate start and stop values. 

*** Please do not change anything to the left of the equals sign. ***

```python
# Set the limits.
start = ...
stop = ...

# Cut the data.
logf_cut = logf_sorted[start:stop]
logG_cut = logG_sorted[start:stop]
errlogG_cut = errlogG_sorted[start:stop]

# Plot the cut data.
fig3 = PHYS231.Scatter(
    xData = logf_cut,
    yData = logG_cut,
    yErrors = errlogG_cut,
    xlabel = "log f",
    ylabel = "log G",
    xUnits = "",
    yUnits = "",
)
```

<!-- END QUESTION -->

## Linear Fits

The final thing that we'll do is fit a straight line through the cut data.  We'll use the 'LinearFit' contained in PHYS231.py to complete this task.  The function call to perform a weighted fit is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(t, y, erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = t, yData = y, yErrors = erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```

For additional information about the 'LinearFit' function, see **Appendix B** near the end of this notebook. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.6</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with appropriate data to complete the weighted linear fit to the high-frequency portion of $\log G$ versus frequency.

*** Please do not change anything to the left of the equals sign. ***

```python
m, b, dm, db, fig = PHYS231.LinearFit(
    xData = ..., 
    yData = ..., 
    yErrors = ..., 
    xlabel = "log f",
    ylabel = "log G",
    xUnits = "",
    yUnits = "",
)
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 2.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell to enter your comments.  This text can be deleted.

<!-- END QUESTION -->

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

***

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
PHYS231.Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
PHYS231.LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS231.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS231
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS231.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: October 29, 2023

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q1.1 (points: None)
```
>>> len(frequency) == len(Vin)
True
```
failure_message: Your frequency and Vin lists are different lengths.
```
>>> len(frequency) == len(errVin)
True
```
failure_message: Your frequency and errVin lists are different lengths.
```
>>> len(frequency) == len(Vout)
True
```
failure_message: Your frequency and Vout lists are different lengths.
```
>>> len(frequency) == len(errVout)
True
```
failure_message: Your frequency and errVout lists are different lengths.
```
>>> len(Vin) == len(errVin)
True
```
failure_message: Your Vin and errVin lists are different lengths.
```
>>> len(Vin) == len(Vout)
True
```
failure_message: Your Vin and Vout lists are different lengths.
```
>>> len(Vin) == len(errVout)
True
```
failure_message: Your Vin and errVout lists are different lengths.
```
>>> len(errVin) == len(Vout)
True
```
failure_message: Your errVin and Vout lists are different lengths.
```
>>> len(errVin) == len(errVout)
True
```
failure_message: Your errVin and errVout lists are different lengths.
```
>>> len(Vout) == len(errVout)
True
```
failure_message: Your Vout and errVout lists are different lengths.


Tests for q1.2 (points: None)
```
>>> len(G) != 0
True
```
failure_message: The G array is empty.
```
>>> all(np.isclose(G, Vout/Vin, atol = 0, rtol = 0.01))
True
```
failure_message: Double check your calculation of G.
```
>>> len(errG) != 0
True
```
failure_message: The errG array is empty.
```
>>> all(np.isclose(errG, np.sqrt((errVout/Vin)**2 + (Vout*errVin/Vin**2)**2), rtol = 0.01))
True
```
failure_message: Double check your calculation of errG.


Tests for q1.4 (points: None)
```
>>> len(errlogG) != 0
True
```
failure_message: The errG array is empty.
```
>>> all(np.isclose(errlogG, (1/np.log(10))*(errG/G), atol = 0, rtol = 0.01))
True
```
failure_message: Double check your calculation of errlogG.
